<a href="https://colab.research.google.com/github/toche7/MLminiDS/blob/main/Lab2_LinearRegression.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Lab 2:  Linear Regression
verison 002

## 1. Example of using SKlearn for model development

https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LinearRegression.html#sklearn.linear_model.LinearRegression


```
from sklearn.linear_model import LinearRegression
reg = LinearRegression()
reg.fit(X, y)
```



| Method | Description |
| ------------|------------ |
| fit(X, y ) | Fit linear model. |
| predict(X) | Predict using the linear model. |
| score(X, y) | Return the coefficient of determination of the prediction.|



### 1.1 Data must be correct format



เนื่องจากการทำนาย X จะเป็นข้อมูลที่ประกอบไปด้วยหลาย feature ได้ นั้นคือ 1 recorded จะประกอบไปด้วยหลายคอลัมน์ได้
ส่วนของ y นั้นจะไปประกอบไปด้วยแค่หนึ่งคอลัมน์

ด้วยเหตุนี้

1. X จะต้องเป็นของมูลระดับ 2 มิติ
2. y จะเป็นของมูลแบบมิติเดียวได้

```python
X = [[0, 0, 1], [1, 1, 2], [2, 2, 3]]
y =  [0, 1, 2]
```

### 1.2 Create model object from a class

ก่อนสร้าง model  จะต้องมีการสร้าง object ที่มาจาก class model ที่เราต้องการเช่น Linear Regression เราก็ต้องสร้าง object ขึ้นมาดังนี้

```python
from sklearn import linear_model
reg = linear_model.LinearRegression()
```

### 1.3 Fit model with train set

เมื่อสร้าง object แล้วเราก็นำไป Training เพื่อหาค่าสัมประสิทธิ์ หรือ weight ของ model ใน library ของ python จะใช้คำว่า fit แทน Training

```python
reg.fit(X,y)
reg.coef_
```

```python
reg.intercept_
```

### 1.4 Apply model with new data

หลังจากผ่านการฝึกแล้วสามารถนำ model ไปทำนายค่าต่างๆ

```python
reg.predict([[0,0,1]])
```

```python
reg.predict(X)
```

### 1.5 Score the model performance

การวัดประสิทธิภาพของ model สามารถทำด้วยการใช้คำสั่ง score ซึ่งจะขึ้นอยู่กับ model ด้วยว่าเป็นโมเดลชนิดไหนการเรียกคำสั่งนี้ก็จะปรับไปตาม model เช่น score ของ Linear Regression ก็จะได้เป็นค่า  $R^2$ ของ model


```python
reg.score(X,y)
```

```python

```

```python

```

```python

```

## 2. Air Plan Fuel Cost Example

### 2.1 Data set

```python
x = [61, 63, 67, 69, 70, 74, 76, 81, 86, 91,95,97]
y = [4.28, 4.08, 4.42, 4.17, 4.48, 4.3, 4.82, 4.7, 5.11, 5.13, 5.64, 5.56]
```

```python
# df = pd.read_excel('https://github.com/toche7/DataSets/raw/main/cost.xlsx')
```

```python
import matplotlib.pyplot as plt

plt.scatter(x,y)
```

### 2.2 Create model object

```python
from sklearn import linear_model
reg1 = linear_model.LinearRegression()
```

### 2.3 Prepare data for model training

```python
import numpy as np
```

```python
X =  np.array(x).reshape(-1,1)
```

```python
X
```

### 2.4 Model training

```python
reg1.fit(X,y)
```

```python
reg1.coef_
```

```python
reg1.intercept_
```

### 2.5 Apply model with some data

```python
reg1.predict(X)
```

```python
yhat = reg1.predict(X)
```

```python
plt.scatter(x,y)
plt.plot(x,yhat,'r');

```

### 2.6 Performance evaluation

```python
reg1.score(X,y)
```

```python

```

## 3. Wildebeest and Zebra

```python
import numpy as py
import pandas as pd
import matplotlib.pyplot as plt
```

### 3.1 ดาวน์โหลด data set



```python
df = pd.read_csv("https://raw.githubusercontent.com/adminho/machine-learning/master/Datasets/predator_dataset.csv", sep=",")
```

ดูข้อมูลใน dataset

```python
#df = pd.read_csv("predator_dataset.csv")
df.head()
```

ลองพล็อตกราฟดู dataset

```python
plt.title("Wildebeest vs Zebra")
plt.xlabel('Wildebeest')
plt.ylabel('Zebra')
plt.plot(df.Wildebeest, df.Zebra, 'ob')
plt.show()
```

### 3.2  Create and train model


```python
from sklearn.linear_model import LinearRegression

regr = LinearRegression()  # create object 'regr' from class LinearRegression
```

```python
 df.Wildebeest.values.reshape(-1,1) # format data for training
```

```python
X = df.Wildebeest.values.reshape(-1, 1)
Y = df.Zebra #.values.reshape(-1, 1)
```

```python
regr.fit(X, Y) # Fitting or Training
```

```python
regr.coef_
```

```python
regr.intercept_
```

### 3.2 Apply Model and Performance Evaluation

แสดงกราฟเส้นตรงที่เกิดจากวิธี regression

```python
from sklearn.metrics import mean_squared_error
predict = regr.predict(X) # prediction
mse = mean_squared_error(Y, predict)
print("intercept=%s, coefficient=%s, MSE=%s" %(regr.intercept_, regr.coef_, mse))
```

```python
plt.title("Wildebeest vs Zebra")
plt.xlabel('Wildebeest')
plt.ylabel('Zebra')
plt.plot(X, Y, 'ob', label='Actual')
plt.plot(X, predict, '-r', label='Predict')
plt.legend(loc='upper left')
plt.show()
```

ข้อมูลไม่เคยพบมาก่อน

```python
x = 200

answer = regr.predict([[x]])
print("มีจำนวนม้าลาย (พันตัว): %.2f" % answer)
```

ลองพล็อตกราฟดูคำตอบ

```python
plt.title("Wildebeest vs Zebra")
plt.xlabel('Wildebeest')
plt.ylabel('Zebra')
plt.plot(X, Y, 'ob', label='Actual')

_X = np.append(X, x)
_predict = np.append(predict, answer)
plt.plot(_X, _predict, '-r', lw=1, label='Predict')
plt.scatter([[x]], answer, s=200, c='g', marker='*', label='answer')
plt.legend(loc='upper left')
plt.show()
```

```python

```

```python

```

```python

```

```python

```

```python

```

## Exercise 1

จากข้อมูลต่อไปนี้จงสร้างแบบจำลองที่เป็น Linear Regression

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

```python
df = pd.read_csv("https://raw.githubusercontent.com/toche7/DataSets/main/sale.txt", header = None)
#df = pd.read_csv("ex1data1.txt", header = None)
df.head()
```

```python
df.columns = ['population','sale']
```

```python
df
```

```python
plt.scatter(df.population, df.sale)
plt.xlabel('Population')
plt.ylabel('Sale')
```

```python
X = df.population
y = df.sale
```

```python
X
```

```python
y
```

#### Ex 1.1 Create Model Object
สร้าง Object ของ Model Linear Regression ใช้ชื่อว่า  regEx1

<!-- grade_id: b56476, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression
regEx1 = None
### BEGIN SOLUTION
regEx1 = LinearRegression()
### END SOLUTION
```

<!-- grade_id: 6149cd, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(isinstance(regEx1, LinearRegression), True)
```

#### Ex 1.2 Train Model
1. เตรียมข้อมูล X และ y เพื่อการ Train model
2. ทำการ Train Model

<!-- grade_id: a16725, points: 0, grade: false, solution: true -->
```python
X = None
### BEGIN SOLUTION
X = df[['population']]
### END SOLUTION

```

<!-- grade_id: 827e20, points: 0, grade: false, solution: true -->
```python
y = None
### BEGIN SOLUTION
y = df.sale
### END SOLUTION

```

```python
regEx1.fit(X,y)
```

```python
regEx1.coef_
```

```python
regEx1.intercept_
```

<!-- grade_id: 3fbe20, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from numpy.testing import assert_almost_equal
assert_almost_equal(regEx1.coef_, np.array([1.19303364]), decimal=4)

```

#### Ex 1.3  Prediction and Score
จงสร้างค่าทำนายและวัดประสิทธิภาพด้วยค่า $R^2$

<!-- grade_id: 4d308d, points: 0, grade: false, solution: true -->
```python
yhat =  None
### BEGIN SOLUTION
yhat = regEx1.predict(X)
### END SOLUTION
```

```python
plt.scatter(X,y)
plt.plot(X,yhat,'r')
```

<!-- grade_id: e69ed9, points: 0, grade: false, solution: true -->
```python
R2 = None
### BEGIN SOLUTION
R2 = regEx1.score(X,y)
### END SOLUTION

```

```python
R2
```

<!-- grade_id: 14685a, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from numpy.testing import assert_almost_equal
assert_almost_equal(R2, 0.7020315537841397, decimal=5)
```

```python

```

```python

```